```python
# Initialize Otter
import otter
grader = otter.Notebook("qbbo.ipynb")
```

```python
import pandas as pd
import numpy as np
%matplotlib inline
import otter
grader = otter.Notebook()

QBBO_PATH = '04222019.QBBO_21_C-decal'
```

# Parsing QBBO

QBBO is NASDAQ's level 1 market data protocol. You can [read its specification online](https://www.nasdaqtrader.com/content/technicalsupport/specifications/dataproducts/QBBOSpecification2.1.pdf). This protocol consists of a repeated series of a message length then a message.

Let's start with some helper functions:

```python
import datetime

def read_int(raw_bytes):
    return int.from_bytes(raw_bytes, byteorder='big')

def read_string(raw_bytes):
    return raw_bytes.decode('ascii')

def parse_timestamp(ns_since_midnight):
    midnight = datetime.datetime(2019, 4, 22)
    delta = datetime.timedelta(seconds=ns_since_midnight / 1e9)
    return midnight + delta
```

**Question 1:** QBBO consists of a series of messages, each consisting of two bytes for length. For example, see the following:

```
00 06 aa bb cc dd ee ff 00 04 14 01 48 03
[len] [    message    ] [len] [ message ]
```

The above contains two messages:
1. `aa bb cc dd ee ff`
2. `14 01 48 03`

A file `04222019.QBBO_21_C-decal` has been provided alongside this assignemnt. It contains real NASDAQ QBBO market data, filtered to only contain system messages and a few tickers (to reduce the size of the download).

Read this file, yielding each message's content.

```python
def read_qbbo_messages(path):
    ...
```

```python
test_reader = read_qbbo_messages(QBBO_PATH)
assert next(test_reader) == b'S\x00\x00\n\x04L`09O'
assert next(test_reader) == b'S\x00\x00\r\x18\xc2\xe5B\xedS'
```

```python
grader.check("q1")
```

**Question 2:** Parse system event messages according to the spec.

```python
def parse_system_event(message):
    ...
    return {
        'type': 'S',
        'timestamp': ...,
        'event': ...
    }
```

```python
assert parse_system_event(b'S\x00\x00\r\x18\xc2\xe5B\xedS') == {
    'type': 'S',
    'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 181),
    'event': 'Start of System Hours',
}
```

```python
grader.check("q2")
```

**Question 3:** Parse quotation messages

```python
def parse_quotation(message):
    ...
    return {
        'type': 'Q',
        'timestamp': ...,
        'ticker': ...,
        'bid_price': ...,
        'bid_size': ...,
        'ask_price': ...,
        'ask_size': ...,
    }
```

```python
assert parse_quotation(b'Q\x00\x00\r\x18\xc3\xe4\xa9GMSFT    Q\x00\x0c5\x00\x00\x00\x00d\x00\x00\x00\x00\x00\x00\x00\x00') == {
    'type': 'Q',
    'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 16919),
    'ticker': 'MSFT',
    'bid_price': 80.0,
    'bid_size': 100,
    'ask_price': 0.0,
    'ask_size': 0
}
```

```python
grader.check("q3")
```

**Question 4:** Parse a message using its message type (the first character). Use the functions you wrote above.

```python
def parse_message(message):
    ...
    
```

```python
assert parse_message(b'S\x00\x00\n\x04L`09O') == {
    'type': 'S',
    'timestamp': datetime.datetime(2019, 4, 22, 3, 3, 33, 577519),
    'event': 'Start of Transmissions',
}
```

```python
assert parse_message(b'Q\x00\x00\r\x18\xc3\xe4\xa9GMSFT    Q\x00\x0c5\x00\x00\x00\x00d\x00\x00\x00\x00\x00\x00\x00\x00') == {
    'type': 'Q',
    'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 16919),
    'ticker': 'MSFT',
    'bid_price': 80.0,
    'bid_size': 100,
    'ask_price': 0.0,
    'ask_size': 0
}
```

```python
grader.check("q4")
```

**Question 5:** Now, write a function to read a QBBO file and yield parsed messages. Use the functions you wrote above.

```python
def read_qbbo_messages_parsed(path):
    ...
```

```python
test_reader = read_qbbo_messages_parsed(QBBO_PATH)
assert next(test_reader) == {
    'type': 'S',
    'timestamp': datetime.datetime(2019, 4, 22, 3, 3, 33, 577519),
    'event': 'Start of Transmissions'
}
assert next(test_reader) == {
    'type': 'S',
    'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 181),
    'event': 'Start of System Hours'
}
assert next(test_reader) == {
    'type': 'Q',
    'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 16919),
    'ticker': 'MSFT',
    'bid_price': 80.0,
    'bid_size': 100,
    'ask_price': 0.0,
    'ask_size': 0
}
```

```python
grader.check("q5")
```

Congratulations! You can now read real marketdata. Although professional trading systems use ITCH (level 2) instead of QBBO (level 1), the protocols are very similar! You can [read ITCH's spec online](https://www.nasdaqtrader.com/content/technicalsupport/specifications/dataproducts/NQTVITCHSpecification.pdf).

Let's do some analysis now that we can read the file.

**Question 6:** Which tickers are included in this file? Write a function that returns a set of ticker names

```python
def tickers_set(path):
    out = set()
    for messsage in read_qbbo_messages_parsed(path):
        if messsage['type'] == 'Q':
            out.add(messsage['ticker'])
    return out
```

```python
assert len(tickers_set(QBBO_PATH)) == 4
```

```python
grader.check("q6")
```

Finally, let's plot the market of `$MSFT`. This part is ungraded:

```python
X_bid = []
Y_bid = []

X_ask = []
Y_ask = []
for message in read_qbbo_messages_parsed(QBBO_PATH):
    if message['timestamp'].hour < 10 or message['timestamp'].hour > 15:
        continue # Skip most out-of-market-hours messages
    if message['type'] == 'Q' and message['ticker'] == 'MSFT':
        if message['bid_price'] > 0:
            X_bid.append(message['timestamp'])
            Y_bid.append(message['bid_price'])
        if message['ask_price'] > 0:
            X_ask.append(message['timestamp'])
            Y_ask.append(message['ask_price'])
        

import matplotlib.pyplot as plt
plt.figure(figsize=(10, 5))
plt.plot(X_ask, Y_ask, label='Ask', color='blue')
plt.plot(X_bid, Y_bid, label='Bid', color='red')
plt.legend()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> test_reader = read_qbbo_messages(QBBO_PATH)
>>> assert next(test_reader) == b'S\x00\x00\n\x04L`09O'
>>> assert next(test_reader) == b'S\x00\x00\r\x18\xc2\xe5B\xedS'
```


Tests for q2 (points: None)
```
>>> assert parse_system_event(b'S\x00\x00\n\x04L`09O') == {'type': 'S', 'timestamp': datetime.datetime(2019, 4, 22, 3, 3, 33, 577519), 'event': 'Start of Transmissions'}
```
```
>>> assert parse_system_event(b'S\x00\x00\r\x18\xc2\xe5B\xedS') == {'type': 'S', 'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 181), 'event': 'Start of System Hours'}
```


Tests for q3 (points: None)
```
>>> assert parse_quotation(b'Q\x00\x00\r\x18\xc3\xe4\xa9GMSFT    Q\x00\x0c5\x00\x00\x00\x00d\x00\x00\x00\x00\x00\x00\x00\x00') == {'type': 'Q', 'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 16919), 'ticker': 'MSFT', 'bid_price': 80.0, 'bid_size': 100, 'ask_price': 0.0, 'ask_size': 0}
```


Tests for q4 (points: None)
```
>>> assert parse_message(b'S\x00\x00\n\x04L`09O') == {'type': 'S', 'timestamp': datetime.datetime(2019, 4, 22, 3, 3, 33, 577519), 'event': 'Start of Transmissions'}
```
```
>>> assert parse_message(b'Q\x00\x00\r\x18\xc3\xe4\xa9GMSFT    Q\x00\x0c5\x00\x00\x00\x00d\x00\x00\x00\x00\x00\x00\x00\x00') == {'type': 'Q', 'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 16919), 'ticker': 'MSFT', 'bid_price': 80.0, 'bid_size': 100, 'ask_price': 0.0, 'ask_size': 0}
```


Tests for q5 (points: None)
```
>>> test_reader = read_qbbo_messages_parsed(QBBO_PATH)
>>> assert next(test_reader) == {'type': 'S', 'timestamp': datetime.datetime(2019, 4, 22, 3, 3, 33, 577519), 'event': 'Start of Transmissions'}
>>> assert next(test_reader) == {'type': 'S', 'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 181), 'event': 'Start of System Hours'}
>>> assert next(test_reader) == {'type': 'Q', 'timestamp': datetime.datetime(2019, 4, 22, 4, 0, 0, 16919), 'ticker': 'MSFT', 'bid_price': 80.0, 'bid_size': 100, 'ask_price': 0.0, 'ask_size': 0}
```


Tests for q6 (points: None)
```
>>> assert len(tickers_set(QBBO_PATH)) == 4
```
